# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays


As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). 

Dictionaries and arrays are both **mutable**, whereas tuples are **immutable**

We'll explain this more below!
Arrays (in my opinion) are one of the most important structures in (Scientific) Computing.

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Output:
```
("penguins", "cats", "sugargliders")
```

```julia
numbers = (1,1,2)
```

Output:
```
(1, 1, 2)
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
@ Base ./client.jl:292

 [14] _start()

    @ Base ./client.jl:495
```

```julia
numbers[1]=5
```

Output:
```
@ Base ./client.jl:292

 [14] _start()

    @ Base ./client.jl:495
```

## Julia Specific: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)
```

A good example is a contacts list, where we associate names with phone numbers.

```julia
myphonebook = Dict("Yiannis" => "0044000000", "Imperial" => "02075895111")
```

Output:
```
Dict{String, String} with 2 entries:
  "Yiannis"  => "0044000000"
  "Imperial" => "02075895111"
```

In this example, each name and number is a "key" and "value" pair. We can grab Yianni's number (a value) using the associated key

```julia
myphonebook["Yiannis"]
```

Output:
```
"0044000000"
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Graduate School"] = "020 7589 5111"
```

Output:
```
"020 7589 5111"
```

Let's check what our phonebook looks like now...

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Yiannis"         => "0044000000"
  "Graduate School" => "020 7589 5111"
  "Imperial"        => "02075895111"
```

We can delete Yiannis from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Yiannis")
```

Output:
```
"0044000000"
```

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Graduate School" => "020 7589 5111"
  "Imperial"        => "02075895111"
```

```julia

```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
myphonebook[1]
```

Output:
```
@ Base ./client.jl:292

 [15] _start()

    @ Base ./client.jl:495
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

## Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```

For an interesting piece of Julia history regarding differences between matrices and vectors feel free to read https://github.com/julialang/julia/issues/4774


For example, we might create an array to keep track of pokemon


```julia
mypokemon = ["Bulbasaur","Squirtle","Charmander","Pikachu"]
```

Output:
```
4-element Vector{String}:
 "Bulbasaur"
 "Squirtle"
 "Charmander"
 "Pikachu"
```

The `1` in `Array{String,1}` means this is a one dimensional vector.  An `Array{String,2}` would be a 2d matrix, etc.  The `String` is the type of each element.

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Yiannis", "Evripides"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Yiannis"
  "Evripides"
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third pokemon listed in `pokemon`, we write

```julia
mypokemon[3]
```

Output:
```
"Charmander"
```

We can use indexing to edit an existing element of an array

```julia
mypokemon[3] = "Mewtwo"
```

Output:
```
"Mewtwo"
```

Yes, Julia is 1-based indexing, not 0-based like Python.  (Although https://docs.julialang.org/en/v1/devdocs/offset-arrays/)

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array. 

Remember the Julia notation ! means the object changes.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
 21
```

and then remove it

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["ice-cream", "chocolate", "eggs"],["penguins", "cats", "spiders"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["ice-cream", "chocolate", "eggs"]
 ["penguins", "cats", "spiders"]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Below are examples of 2D and 3D arrays populated with random values.

```julia
rand(4, 3)
```

Output:
```
4×3 Matrix{Float64}:
 0.0407491  0.0925724  0.708855
 0.0397694  0.93413    0.236814
 0.312578   0.213933   0.787578
 0.777425   0.785905   0.325041
```

```julia
rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.0141338  0.0957528  0.317101
 0.360828   0.0508514  0.484834
 0.0286108  0.518769   0.471675
 0.0746651  0.962424   0.526508

[:, :, 2] =
 0.957841  0.771966  0.807781
 0.773477  0.304807  0.511314
 0.723581  0.467766  0.328393
 0.562373  0.658319  0.883337
```

```julia
?rand
```

Be careful when you want to copy arrays!

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers = fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
 404
   1
   2
   3
   5
   8
  13
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somemorenumbers = copy(fibonacci)
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somemorenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

# Vectors and Matrices
We can also create vectors : 

```julia
squares = [1, 4, 9, 16, 25, 36, 49, 64,81]
```

Output:
```
9-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
 81
```

```julia
squares[1]
```

Output:
```
1
```

```julia
squares[1:3]
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

```julia
squares[end]
```

Output:
```
81
```

```julia
squares[-1]
```

Output:
```
@ Base ./client.jl:292

 [15] _start()

    @ Base ./client.jl:495
```

The below code will give us a first glimpse into the powerful **type** system present in Julia

```julia
typeof(squares) 
```

Output:
```
Vector{Int64} (alias for Array{Int64, 1})
```

# Concatenation
If, instead of commas, we use spaces, then the values are concatenated horizontally

```julia
cubes = [1, 8, 27, 64, 125, 216, 343, 512]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
cubes2 = [1 8 27]
```

Output:
```
1×3 Matrix{Int64}:
 1  8  27
```

```julia
pop!(squares)
```

Output:
```
81
```

```julia
squares
```

Output:
```
8-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
```

```julia
a = [1:3 4:6]
b = [1:3; 4:6]
```

Output:
```
6-element Vector{Int64}:
 1
 2
 3
 4
 5
 6
```

```julia
powers = [1:8 squares cubes]
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

```julia
powers[4, 2]
```

Output:
```
16
```

```julia
powers[:, 3]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
powers[7, :]
```

Output:
```
3-element Vector{Int64}:
   7
  49
 343
```

```julia
typeof(powers)
```

Output:
```
Matrix{Int64} (alias for Array{Int64, 2})
```

Semicolon separators perform vertical concatenation:

```julia
[squares; cubes]
```

Output:
```
16-element Vector{Int64}:
   1
   4
   9
  16
  25
  36
  49
  64
   1
   8
  27
  64
 125
 216
 343
 512
```

Whereas commas would simply create an array of arrays : 

```julia
# dimension of vectors (row/column) is automatically determined
nested_powers = [[1,2,3,4,5,6,7,8], squares, cubes]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3, 4, 5, 6, 7, 8]
 [1, 4, 9, 16, 25, 36, 49, 64]
 [1, 8, 27, 64, 125, 216, 343, 512]
```

```julia
nested_powers[2]
```

Output:
```
8-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
```

Horizontal and vertical concatenation can be used together to as a simple syntax for matrix literals

```julia
[1 3 5; 2 4 6]
```

Output:
```
2×3 Matrix{Int64}:
 1  3  5
 2  4  6
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```



Be careful of 1-based indexing! What do we think the following will produce?

```julia
for i in 1:10 #python 1,,,9
    println(i)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

We could use a for loop to generate the same results as either of the examples above:

```julia
A = fill(0, (8, 3)) # Allocate an 8x3 matrix to store the values into
for pow in 1:3
    for value in 1:8
        A[value, pow] = value ^ pow
    end
end
A
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

```julia
A == powers
```

Output:
```
true
```


## Array Comprehensions

```julia
squares = [value^2 for value in 1:8]
```

Output:
```
8-element Vector{Int64}:
  1
  4
  9
 16
 25
 36
 49
 64
```

```julia
cubes = [value^3 for value in 1:8]
```

Output:
```
8-element Vector{Int64}:
   1
   8
  27
  64
 125
 216
 343
 512
```

```julia
powers = [value^pow for value in 1:8, pow in 1:3]
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

# The element type

Note that every time an array prints out, it is displaying its element type and dimensionality, for example `Array{Int64, 2}`. This describes what it can store — and thus what it can return upon indexing.

```julia
powers
```

Output:
```
8×3 Matrix{Int64}:
 1   1    1
 2   4    8
 3   9   27
 4  16   64
 5  25  125
 6  36  216
 7  49  343
 8  64  512
```

```julia
typeof(powers)
```

Output:
```
Matrix{Int64} (alias for Array{Int64, 2})
```

```julia
typeof(powers[1, 1])
```

Output:
```
Int64
```

Further, the array will try to convert any new values assigned into it to its element type:

```julia
powers[1, 1] = 1.6
```

Output:
```
@ Base ./client.jl:292

 [17] _start()

    @ Base ./client.jl:495
```

```julia
powers[1, 1] = -5.0 # This can be losslessly converted to an integer
```

Output:
```
-5.0
```

```julia
powers
```

Output:
```
8×3 Matrix{Int64}:
 -5   1    1
  2   4    8
  3   9   27
  4  16   64
  5  25  125
  6  36  216
  7  49  343
  8  64  512
```

Arrays that have an exact and concrete element type are generally significantly faster, so Julia will try to find an amenable element type for you in its literal construction syntax:

```julia
fortytwosarray = [42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

Output:
```
6-element Vector{Float64}:
 42.0
 42.0
 42.0
 42.0
 42.0
 42.0
```

```julia
for x in fortytwosarray
    show(x)
    println("\tisa $(typeof(x))")
end
```

Output:
```
42.0	isa Float64
42.0	isa Float64
42.0	isa Float64
42.0	isa Float64
42.0	isa Float64
42.0	isa Float64
```

The `Any` array can be helpful for disabling these behaviors and allowing all kinds of different objects:

```julia
anyfortytwos = Any[42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a]
```

Output:
```
6-element Vector{Any}:
   42
   42.0
   42.0
   42.0f0
  42//1
 0x2a
```

```julia
anyfortytwos[1] = "FORTY TWO"
anyfortytwos
```

Output:
```
6-element Vector{Any}:
     "FORTY TWO"
   42.0
   42.0
   42.0f0
  42//1
 0x2a
```

### Exercises

#### 4.1 
Create an array, `a_ray`, with the following code:

```julia
a_ray = [1, 2, 3]
```

Add the number `4` to the end of this array and then remove it.

```julia
a_ray = [1, 2, 3]
push!(a_ray, 4)
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 4.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Emergency"       => "911"
  "Graduate School" => "020 7589 5111"
  "Imperial"        => "02075895111"
```

```julia
myphonebook["Emergency"] = "911"
```

Output:
```
"911"
```

#### 4.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 4.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia
flexible_phonebook["Emergency"] = 911
```

Output:
```
911
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 4.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values? (hint: try using [Julia's documentation for dictionaries](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

#### 4.6 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [x ^ 2 for x = 1 : 100]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
     ⋮
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

```julia

```